---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Maximum Likelihood Estimate (MLE)

```yaml
type: VideoExercise
key: 8f8cb74702
xp: 50
```

`@projector_key`
07185d5d9a25552df760cdb9f65fef3f
